```python
# Initialize Otter
import otter
grader = otter.Notebook("Linear transforms.ipynb")
```

# Matrix Analysis 2024 - EE312
## Week 2  - Linear transforms
[N. Aspert](https://people.epfl.ch/nicolas.aspert) - [LTS2](https://lts2.epfl.ch)

The first week notebook (introduction to Python, Numpy and Matplotlib) can be used as a help.

## Important
You need to submit *individually* your answers on moodle before the next exercise session. For the theoretical questions you can either fill the notebook or write it on a separate sheet (if you are not comfortable with Markdown/TeX) and upload a scanned version. 

## Objective

The end goal is to understand purely algebraic, matrix based, view of a few linear transforms. You will use those linear transform to perform some basic time-frequency analysis of signals.

```python
import numpy as np
import matplotlib.pyplot as plt
```

## Part I - Fourier

<!-- BEGIN QUESTION -->

1. Prove that any set of orthogonal vectors $v_i \in \mathbb{C}^N, \, i=1, \ldots , M \leq N$ such that $v_i^H v_j = C \delta_{i,j}$ is linearly independent (where $C$ is some constant).

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

2. Compute $a_r = \sum_{n=0}^{N-1}e^{j2\pi r\frac{n}{N}}$, where $r$ is an integer (discuss the result depending on the value of $r$).

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

3. Let $v_k \in \mathbb{C}^N$ be such that $v_k[n] = e^{-j 2 \pi \frac{kn}{N}}$, for $k,n = 0, \ldots , N-1$. 
- Prove that these vectors are mutually orthogonal, hence linearly independent. 
- Compute the norm of $v_k$.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

4. Implement the function `get_fourier_matrix` that returns a normalized Fourier matrix of size $N\times N$. Do not make use of the builtin DFT/FFT functions in `numpy` or `scipy`. Raise a `ValueError` exception if a ngetive or zero $N$ value is supplied.

```python
def get_fourier_matrix(N):
    ...
```

```python
grader.check("q4")
```

Let us now generate two test signals. 

The first one $x_1(t)$ is made of four piecewise sinusoids, of different frequencies:

$$
x_1(t) = \cos(2\pi 5t), 0\le t < 2\\
x_1(t) = \cos(2\pi 10t), 2\le t < 4\\
x_1(t) = \cos(2\pi 25t), 4\le t < 6\\
x_1(t) = \cos(2\pi 50t), 6\le t < 8\\
$$


```python
Fs = 256 # sampling frequency
t = np.arange(0, Fs*8)/Fs
x1 = np.zeros(Fs*8)
x1[0:Fs*2] = np.cos(2*np.pi*5*t[0:Fs*2])
x1[Fs*2:Fs*4] = np.cos(2*np.pi*10*t[Fs*2:Fs*4])
x1[Fs*4:Fs*6] = np.cos(2*np.pi*25*t[Fs*4:Fs*6])
x1[Fs*6:Fs*8] = np.cos(2*np.pi*50*t[Fs*6:Fs*8])
```

The second signal $x_2(t)$ is the sum of the same sinusoids over the complete time interval, with a scaling term s.t. the amplitude of both signals is identical.

```python
x2 = 0.25*(np.cos(2*np.pi*5*t) + np.cos(2*np.pi*10*t) + np.cos(2*np.pi*25*t) + np.cos(2*np.pi*50*t))
```

<!-- BEGIN QUESTION -->

5. 
- Display the generated signals using `plt.plot`. 
- Compute their Fourier transforms using the Fourier matrix generate at the previous question.
- Display the amplitude of their Fourier spectrum. What do you observe ? 

_Type your answer here, replacing this text._

```python
# plot x1
...
```

```python
# plot x2
...
```

```python
# Compute the Fourier transform of x1 and x2
...
```

```python
# Plot the spectrum of x1
...
```

```python
# Plot the spectrum of x2
...
```

<!-- END QUESTION -->

In order to have a better compromise between time and frequency, the input signal will be split in smaller non-overlapping blocks of length $p$, and we will perform the DFT of each block.

6. Using the `get_fourier_matrix` implemented previously, fill the `get_block_dft_matrix` function below s.t. it returns a $N\times N$ matrix that will perform the block Fourier transform when applied to an input vector. Raise a `ValueError` if $p$ does not divide $N$.

Hint: [numpy.pad](https://numpy.org/doc/stable/reference/generated/numpy.pad.html#numpy.pad) and/or [numpy.kron](https://numpy.org/doc/stable/reference/generated/numpy.kron.html) might prove useful.

```python
def get_block_dft_matrix(N, p):
    ...
```

```python
grader.check("q6")
```

<!-- BEGIN QUESTION -->

We will now use this block Fourier transform to how the frequencies of the signal evolve through time.

7. 
- Using the `reshape` and `plt.imshow` functions, display the amplitude of the result of the block Fourier transform of $x_1$ and $x_2$. Is the result improved when compared to the one observed in question 5 ?
- What is the influence of parameter $p$ ? 


_Type your answer here, replacing this text._

```python
# Compute the block DFT matrix Wb
...
```

```python
# Plot the block DFT of x1
...
```

```python
# Plot the block DFT of x2
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

8. In a real-world application, would generating a $N\times N$ matrix to perform the block Fourier transform be a good way to implement it ? (Justify)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Part II - Haar transform

In this part we will study another approach to study the time/frequency properties of signals. 

Let us consider a vector $x\in\mathbb{R}^N$, with $N$ being even. The single-level Haar transform decomposes $x$ into two vectors $a^1$ and $d^1$ belonging to $\mathbb{R}^{\frac{N}{2}}$. 

The coefficients of $a^1$ are defined as follows: $a^1_n = \frac{1}{\sqrt{2}}(x_{2n-1} + x_{2n}), n=1, ..., \frac{N}{2}$. $a^1$ is referred to as the *average coefficients vector*.

The coefficients of $d^1$ are defined as follows: $d^1_n = \frac{1}{\sqrt{2}}(x_{2n-1} - x_{2n}), n=1, ..., \frac{N}{2}$. $d^1$ is referred to as the *detail coefficients vector*.


<!-- BEGIN QUESTION -->

9. Let us represent the single-level Haar transform by a matrix $H_1$ s.t. 
$$
H_1 x = \begin{pmatrix}a^1 \\ d^1\end{pmatrix}
$$
Prove that $H_1$ is orthonormal.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

10. Write a function that returns the single-level Haar transform matrix $H_1$ for a given $N$. Raise a `ValueError` if $N$ is invalid.

```python
def get_sl_haar_matrix(N):
    ...
```

```python
grader.check("q10")
```

The multi-level Haar transform is defined by recursively applying the single-level transform **to the average coefficient parts**. 

For instance constructing 2-level Haar transform over $N$ points start with the previously defined $H_{1,N}$ matrix of size $N\times N$ and the corresponding $\frac{N}{2}\times\frac{N}{2}$ version denoted by $H_{1,\frac{N}{2}}$. 

$H_{1,N}$ can be written as
$$
\begin{pmatrix} H_{1, N}^a \\ H_{1,N}^d \end{pmatrix},
$$
where $H_{1, N}^a$ and $H_{1,N}^d$ are respectively the average and detail coefficient matrices, both of size $\frac{N}{2}\times N$.

Following these notations, the 2-level Haar transform matrix $H_{2,N}$ can be written as:
$$
\begin{pmatrix} H_{1,\frac{N}{2}}H_{1, N}^a \\ H_{1,N}^d \end{pmatrix},
$$

11. Implement a function that returns the $H_{p,N}$ matrix of size $N\times N$ that performs a $p$-level haar transform. Raise a `ValueError` if the size and the level are incompatible, or if the level is smaller than 1.

```python
def get_haar_matrix(N, level):
    ...
```

```python
grader.check("q11")
```

<!-- BEGIN QUESTION -->

12. Prove that $H_{p,N}$ is orthonormal.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Haar transform visualization

In order to make the visualization of the Haar decomposition easy, we provide you the `plot_haar_coeffs` function that will display the average and detail coefficients of the different levels. 

The function takes 2 arguments:
- the input signal
- the number of levels


<!-- BEGIN QUESTION -->

13. Display the Haar transform of $x1$ and $x2$.

```python
from display_helper import plot_haar_coeffs
```

```python
# display the decomposition of x1
...
```

```python
# display the decomposition of x2
...
```

<!-- END QUESTION -->

## Part III - Denoising

We will now use the different transforms defined in part I and II to perform denoising.
Let us create a noisy signal for this purpose.

```python
angle1 = np.linspace(0, 5*np.pi/2, 300)
wave1  = np.sin(angle1)
angle2 = np.linspace(0, 3*np.pi/2, 300)
wave2  = np.sin(angle2)
angle3 = np.linspace(np.pi/2, 2*np.pi, 424)
wave3  = np.sin(angle3)
wave   = np.append(wave1, wave2)
wave   = np.append(wave, wave3)
wave_noisy = wave + 0.2*np.random.normal(0, 1, 1024)
```

```python
plt.plot(wave_noisy, 'r')
plt.plot(wave)
```

The noise is usually located in the higher frequencies. However, the signal we created is a bit special as it has two discontinuities which also generate high frequencies components (remember the Fourier transform of a rectangle function is a sinc). 

<!-- BEGIN QUESTION -->

14. Implement a function `denoise_signal` that perform denoising of the input signal by using a supplied orthonormal transform matrix, and by setting the transformed coefficients having an amplitude smaller than a given threshold to 0. You might want to use the [numpy.where](https://numpy.org/doc/stable/reference/generated/numpy.where.html) function for this. When denoising using the Haar transform, you can preform the thresholding only on the detail coefficients. Implement the function `denoise_signal_haar` that performs this operation.

NB: The result returned should be real, in order to be displayed. 

```python
def denoise_signal(W, x, threshold=0.1):
    """
    W: NxN input orthonormal matrix (Fourier, block-Fourier or Haar)
    x: input signal (of length N)
    """
    ...
```

```python
def denoise_signal_haar(W, x, threshold=0.1, detail_start_index=256):
    """
    W: NxN input orthonormal matrix (Fourier, block-Fourier or Haar)
    x: input signal (of length N)
    detail_start_index: thresholding is performed on x[detail_start_index:]
    """
    ...
```

```python
# Perform denoising with the full Fourier transform and display the result. 
# Make sure you choose a good threshold
...
```

```python
# Perform denoising with the block Fourier transform and display the result
# Make sure you choose a good threshold and block size
...
```

```python
# Perform denoising with the Haar transform and display the result
# Make sure you choose a good threshold and an appropriate number of levels
...
```

```python
grader.check("q14")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

15. Compare the three denoising methods (Fourier, block Fourier and Haar). Which one performs better (in terms of noise removal but also in terms of discontinuity preservation) ? Was that expected (justify) ?

_Type your answer here, replacing this text._

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Upload your notebook and separate pdf for theoretical questions if needed

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q10 (points: 4)
```
>>> with np.testing.assert_raises(ValueError):
...     get_sl_haar_matrix(-1)
>>> with np.testing.assert_raises(ValueError):
...     get_sl_haar_matrix(0)
>>> with np.testing.assert_raises(ValueError):
...     get_sl_haar_matrix(3)
```
failure_message: Did you forget to validate the input size before performing the computation ?
success_message: Good, you properly validated size before computing the result
```
>>> np.testing.assert_array_almost_equal(get_sl_haar_matrix(2), np.array([[1, 1], [1, -1]]) / np.sqrt(2))
>>> np.testing.assert_array_almost_equal(get_sl_haar_matrix(4), np.array([[1, 1, 0, 0], [0, 0, 1, 1], [1, -1, 0, 0], [0, 0, 1, -1]]) / np.sqrt(2))
```
failure_message: Results seem incorrect, check your implementation
success_message: Good, results look correct
```
>>> H = get_sl_haar_matrix(16)
>>> np.testing.assert_array_almost_equal(H @ H.T, np.eye(16))
>>> np.testing.assert_array_almost_equal(H.T @ H, np.eye(16))
```
failure_message: Results seem incorrect, check your implementation
success_message: Good, computed matrix is orthogonal


Tests for q11 (points: 4)
```
>>> with np.testing.assert_raises(ValueError):
...     get_haar_matrix(-1, 0)
>>> with np.testing.assert_raises(ValueError):
...     get_haar_matrix(16, 0)
>>> with np.testing.assert_raises(ValueError):
...     get_haar_matrix(4, 3)
```
failure_message: Did you forget to validate the input size before performing the computation ?
success_message: Good, you properly validated size/level before computing the result
```
>>> H = get_haar_matrix(4, 2)
>>> a = 1 / np.sqrt(2)
>>> np.testing.assert_array_almost_equal(H, [[0.5, 0.5, 0.5, 0.5], [0.5, 0.5, -0.5, -0.5], [a, -a, 0.0, -0.0], [0.0, -0.0, a, -a]])
```
failure_message: Results seem incorrect, check your implementation
success_message: Good, results look correct
```
>>> H = get_haar_matrix(64, 4)
>>> np.testing.assert_array_almost_equal(H @ H.T, np.eye(64))
>>> np.testing.assert_array_almost_equal(H.T @ H, np.eye(64))
```
failure_message: Results seem incorrect, check your implementation
success_message: Good, matrix is orthogonal


Tests for q14 (points: 6)
```
>>> waverec = denoise_signal(get_fourier_matrix(len(wave_noisy)), wave_noisy, 0.0)
>>> np.testing.assert_array_almost_equal(waverec, wave_noisy)
```
failure_message: Results seem incorrect, check your implementation
success_message: Good, not denoising the noisy signal does not modify it.


Tests for q4 (points: 1)
```
>>> from scipy.linalg import dft
>>> np.testing.assert_array_almost_equal(get_fourier_matrix(16), dft(16, scale='sqrtn'))
```
failure_message: Check your implementation
success_message: Good, your implementation returns correct results
```
>>> from unittest.mock import patch
>>> 
>>> def check_fourier_dft(N):
...     with patch('scipy.linalg.dft') as mock_dft:
...         get_fourier_matrix(N)
...         mock_dft.assert_not_called()
>>> check_fourier_dft(32)
```
failure_message: Do not use scipy.linalg.dft in your implementation
success_message: Good, you did not use the scipy.linalg.dft function
```
>>> from unittest.mock import patch
>>> 
>>> def check_fourier_fft(N):
...     with patch('scipy.fft.fft') as mock_fft:
...         get_fourier_matrix(N)
...         mock_fft.assert_not_called()
>>> check_fourier_fft(24)
```
failure_message: Do not use scipy.fft.fft in your implementation
success_message: Good, you did not use the scipy.fft.fft function
```
>>> from unittest.mock import patch
>>> 
>>> def check_fourier_npfft(N):
...     with patch('numpy.fft.fft') as mock_fft:
...         get_fourier_matrix(N)
...         mock_fft.assert_not_called()
>>> check_fourier_npfft(48)
```
failure_message: Do not use numpy.fft.fft in your implementation
success_message: Good, you did not use the numpy.fft.fft function
```
>>> W = get_fourier_matrix(8)
>>> np.testing.assert_array_almost_equal(np.conj(W.T) @ W, np.eye(8))
```
failure_message: Check your implementation, resulting matrix should be orthogonal.
success_message: Good, your implementation returns an orthogonal matrix
```
>>> with np.testing.assert_raises(ValueError):
...     get_fourier_matrix(0)
>>> with np.testing.assert_raises(ValueError):
...     get_fourier_matrix(-5)
```
failure_message: Did you forget to validate the input size before performing the computation ?
success_message: Good, you properly validated size before computing the result


Tests for q6 (points: 4)
```
>>> W = get_block_dft_matrix(32, 8)
>>> np.testing.assert_array_almost_equal(np.conj(W.T) @ W, np.eye(32))
```
failure_message: Check your implementation, resulting matrix should be orthogonal.
success_message: Good, your implementation returns an orthogonal matrix
```
>>> with np.testing.assert_raises(ValueError):
...     get_block_dft_matrix(64, 7)
```
failure_message: Did you forget to validate the input size before performing the computation ?
success_message: Good, you properly validated size before computing the result
```
>>> W = get_block_dft_matrix(4, 2)
>>> np.testing.assert_array_almost_equal(W[0:2, 0:2], get_fourier_matrix(2))
>>> np.testing.assert_array_almost_equal(W[2:4, 2:4], get_fourier_matrix(2))
>>> np.testing.assert_array_almost_equal(W[2:4, 0:2], np.zeros((2, 2)) * 1j)
>>> np.testing.assert_array_almost_equal(W[0:2, 2:4], np.zeros((2, 2)) * 1j)
```
failure_message: Check your implementation !
success_message: Good, your implementation returns the correct result
