```python
# Initialize Otter
import otter
grader = otter.Notebook("Registration.ipynb")
```

# Matrix Analysis 2024 - EE312
## Week 11 - Rigid registration and SVD
[LTS2](https://lts2.epfl.ch)


In this session you will study how to solve a rigid registration problem (in the least square sense) using the SVD.

Let us consider two sets of points from $\mathbb{R}^d$, $\mathcal{P}=\{\mathbf{p}_k\}$ and $\mathcal{Q}=\{\mathbf{q}_k\}$, $k=1,...,n$. Each point $\mathbf{p}_k$ matches a point $\mathbf{q}_k$ through a linear transformation, i.e. $\mathbf{q}_k = \mathcal{T}(\mathbf{p}_k)$. In this exercise, we will consider the transformations that are a combination of a translation $\mathbf{t}$ (represented by a vector in $\mathbb{R}^d$) and a rotation $R$ (represented by an orthogonal $d\times d$ matrix).

Rigid registration consists in finding the optimal (in the least-squares sense) $R$ and $t$ such that
$$
(R,\mathbf{t}) = \underset{t\in\mathbb{R}^d,R\in SO(d)}{\operatorname{argmin}}\sum_{k=1}^n \| (R\mathbf{p}_k+\mathbf{t}) - \mathbf{q}_k \|^2
$$

## 1. Finding the rotation matrix $R$
Let us first consider the (simpler) case where the translation component of the transformation is null, and we only have a rotation $R$ to find, i.e.

$$
R = \underset{R\in SO(d)}{\operatorname{argmin}}\sum_{k=1}^n \| R\mathbf{p}_k - \mathbf{q}_k \|^2
$$

<!-- BEGIN QUESTION -->

1.1 Show that finding $R$ is equivalent to maximize the quantity $\sum_{k=1}^n \mathbf{q}_k^T R \mathbf{p}_k$.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

Let us define the $d\times n$ matrices $P$ and $Q$ made from the column vectors $\mathbf{p}_k$ and $\mathbf{q}_k$, i.e.:
$$
P = \begin{pmatrix}\mathbf{p}_1  & \mathbf{p}_2 & ... & \mathbf{p}_n\end{pmatrix}, Q = \begin{pmatrix}\mathbf{q}_1  & \mathbf{q}_2 & ... & \mathbf{q}_n\end{pmatrix},
$$

<!-- BEGIN QUESTION -->

1.2 Prove that $\sum_{k=1}^n \mathbf{q}_k^T R \mathbf{p}_k = \mathrm{tr}(Q^TRP)$

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

1.3 Let us introduce the matrix $S=PQ^T$, its SVD decomposition $S=U\Sigma V^T$ and $M=V^TRU$, show that $\mathrm{tr}(Q^TRP)=\mathrm{tr}(\Sigma M)$.


*Reminder*: for matrices of compatible size, the following equality holds: $\mathrm{tr}(AB)=\mathrm{tr}(BA)$.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

1.4 Let us denote by $m_{ij}$ the coefficients of $M$. Show that $|m_{ij}|\leq 1$. Finally show that $\mathrm{tr}(\Sigma M)$ is maximum for $M=I$ and compute the corresponding value of $R$

(Hint: prove that $M$ is orthogonal first).

_Type your answer here, replacing this text._

<!-- END QUESTION -->

1.5 We can now verify that our method works using a numerical example, using $d=2$ to simplify visualization.

```python
import numpy as np
import matplotlib.pyplot as plt
```

```python
# Load data points from text file
fish = np.loadtxt('./data/fish.txt')
```

```python
plt.scatter(fish[0, :], fish[1, :], marker='+')
```

```python
# Generate a rotated version
theta = -np.pi/6 # you can change this
R = np.eye(2) # replace this with a rotation matrix depending on theta
Y = R@fish  + 0.05*np.random.randn(fish.shape[1]) # transform + add noise
```

```python
plt.scatter(fish[0, :], fish[1, :], marker='+')
plt.scatter(Y[0, :], Y[1, :], marker='x')
```

```python
def compute_R(P, Q):
    """
    Compute the rotation matrix that transforms P into Q

    Parameters
    ----------
    P : input data matrix (2xN)
    Q: transformed data matrix (2xN)
    
    Returns
    -------
    the estimated rotation matrix
    """
    ...
```

```python
R_r = compute_R(fish, Y)
yr = R@fish
```

```python
plt.scatter(yr[0, :], yr[1, :], marker='+')
plt.scatter(Y[0, :], Y[1, :], marker='x')
```

```python
grader.check("q1p5")
```

## 2. Finding the translation $\mathbf{t}$


<!-- BEGIN QUESTION -->

2.1 Let us consider the function 
$$
F(\mathbf{t}) = \sum_{k=1}^n \| (R\mathbf{p}_k+\mathbf{t}) - \mathbf{q}_k \|^2,
$$ with $\mathbf{t}\in\mathbb{R}^d$.

Let us also introduce the following quantities $\bar{\mathbf{p}}=\frac{1}{n}\sum_{k=1}^n\mathbf{p}_k$ and $\bar{\mathbf{q}}=\frac{1}{n}\sum_{k=1}^n\mathbf{q}_k$.

Solve
$$
\frac{\partial F(\mathbf{t})}{\partial \mathbf{t}} = 0
$$ 
and find the optimal value of $\mathbf{t}$

Hint: if $x\in\mathbb{R}^d$ is a vector, the following holds:
$$
\frac{\partial x^Tx}{\partial x} = 2x^T
$$

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

2.2 Let us define $\mathbf{p}^\prime_k = \mathbf{p}_k - \bar{\mathbf{p}}$ and $\mathbf{q}^\prime_k = \mathbf{q}_k - \bar{\mathbf{q}}$. Prove that 
$$
F(\mathbf{t}) = \sum_{k=1}^n \| R\mathbf{p}^\prime_k - \mathbf{q}^\prime_k \|^2,
$$

_Type your answer here, replacing this text._

<!-- END QUESTION -->

2.3 Now we have found both $R$ and $\mathbf{t}$, we can test on a numerical example

```python
# Generate a rotated + translated version
t = np.array([-0.5, 1]).reshape(2,1) # change this if you wish
Yt = R@fish  + t + 0.05*np.random.randn(fish.shape[1]) # transform + add noise
```

```python
plt.scatter(fish[0, :], fish[1, :], marker='+')
plt.scatter(Yt[0, :], Yt[1, :], marker='x')
```

```python
def do_registration(P, Q):
    """
    Performs the registration (translation + rotation) that transforms P into Q

    Parameters
    ----------
    P : input data matrix (2xN)
    Q: transformed data matrix (2xN)
    
    Returns
    -------
    the translated/rotated P that fits Q
    """
    ...
```

```python
Yr = do_registration(fish, Yt)
```

```python
plt.scatter(Yt[0, :], Yt[1, :], marker='+')
plt.scatter(Yr[0, :], Yr[1, :], marker='x')
```

```python
grader.check("q2p3")
```

Tests for q1p5 (points: None)
```
>>> Q = np.array([fish[1, :], fish[0, :]])
>>> P = fish
>>> R_o = compute_R(P, Q)
>>> np.testing.assert_array_almost_equal(R_o, [[0, 1], [1, 0]])
```
failure_message: Check your implementation
success_message: Good, rotation estimation looks correct


Tests for q2p3 (points: None)
```
>>> t = np.array([0.5, -1]).reshape(2, 1)
>>> R = np.array([[0, 1], [1, 0]])
>>> Yt = R @ fish + t
>>> P = fish
>>> np.testing.assert_array_almost_equal(do_registration(P, Yt), Yt)
```
failure_message: Check your implementation
success_message: Good, registration estimation looks correct
